# Working with [ICD9](https://www.cdc.gov/nchs/icd/icd9.htm) Data

```python
%matplotlib inline
```

```python
import pickle
from collections import defaultdict
from nose.tools import assert_equals
```

```python
with open("icd9.pickle","rb") as f0:
    icd9_data = pickle.load(f0)
```

## What does our data look like?

```python
icd9_data[:5]
```

```python
for d in icd9_data:
    try:
        if "HEP" in d[2]:
            print(d)
    except Exception as Error:
        pass
```

### Each element of `icd9_data` is a tuple with three elements
1. A patient id
1. An ICD9 code
1. The label for that ICD9 code

#### Create a dictionary named `icd9_map` with keys ICD9 code and values the ICD9 label

strip extra white spaces

<!-- grade_id: cell-5d44913c5eba6a2c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
icd9_patients = defaultdict(list)
patient_diagnoses = defaultdict(list)
```

### How many patients are there for each diagnosis?

Create a list named `icd9_patients_list` sorted by the number of patients per diagnosis.

<!-- grade_id: cell-ae2a3dc827383e7f, points: 0, grade: false, solution: true -->
```python
icd9_patients = defaultdict(list)
# YOUR CODE HERE
raise NotImplementedError()
```

```python
import random
d,p = random.choice(icd9_patients_list)
print(d,len(p),sep="\n")
print(icd9_patients["V12.59"])
for d,p in icd9_patients.items():
    if "HX" in d:
        print(d,len(p))
```

```python
assert_equals(len(icd9_patients["V12.59"]),5)
assert_equals(len(icd9_patients["572.2"]),12)
```

#### Loop through  `icd9_patients_list`
1. for each element in `icd9_patients_list` print the ICD9 label corresponding to the code and the number of patients with each diagnosis.

<!-- grade_id: cell-121dd04abe83faee, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### How many diagnoses does each patient have?

<!-- grade_id: cell-305cd7b053c0116a, points: 0, grade: false, solution: true -->
```python
patient_diagnoses = defaultdict(list)
# YOUR CODE HERE
raise NotImplementedError()
```

```python
for p, d in patient_diagnoses_list:
    print(p)
    print(d) # replace with len(d)
    print("\n")
```

<!-- grade_id: cell-322e571e6bd058c0, points: 0, grade: true, solution: false -->
```python
assert_equals(len(patient_diagnoses[2512]),49)
assert_equals(len(patient_diagnoses[353]),56)
assert_equals(len(patient_diagnoses_list[0][1]),125)
assert_equals(len(patient_diagnoses_list[45][1]),23)

```

```python

```