# Working with [ICD9](https://www.cdc.gov/nchs/icd/icd9.htm) Data

```python
%matplotlib inline
```

```python
import pickle
from collections import defaultdict
from nose.tools import assert_equals
```

```python
with open("icd9.pickle","rb") as f0:
    icd9_data = pickle.load(f0)
```

## What does our data look like?

```python
icd9_data[:5]
```

### Each element of `icd9_data` is a tuple with three elements
1. A patient id
1. An ICD9 code
1. The label for that ICD9 code

```python
for d in icd9_data:
        if "HEP" in d[2]:
            print(d)

```

```python
d
```

```python
for d in icd9_data:
    try:
        if "HEP" in d[2]:
            print(d)
    except Exception as Error:
        pass
```

#### Create a dictionary named `icd9_map` with keys ICD9 code and values the ICD9 label

strip extra white spaces

<!-- grade_id: cell-5d44913c5eba6a2c, points: 0, grade: false, solution: true -->
```python
icd9_data
icd9_map = {}
for d in icd9_data:
    key = d[1]
    try:
        value = d[2].strip()
    except:
        value = "NO LABEL PROVIDED"
    icd9_map[key]=value
```

```python
icd9_map
```

```python
icd9_patients = defaultdict(list)
patient_diagnoses = defaultdict(list)
```

### How many patients are there for each diagnosis?

Create a list named `icd9_patients_list` sorted by the number of patients per diagnosis.

### Use a default dictionary to aggregate the data

<!-- grade_id: cell-ae2a3dc827383e7f, points: 0, grade: false, solution: true -->
```python
icd9_patients = defaultdict(list)
for p,c,l in icd9_data:
    try:
        icd9_patients[c].append(p)
    except:
        pass
print(icd9_patients)    


#raise NotImplementedError()
```

```python
icd9_patients_list=list(icd9_patients.items())
icd9_patients_list.sort()
print(icd9_patients_list)
```

```python
icd9_patients_list.sort()
```

```python
def length_2_elem(x):
    return len(x[1])

```

```python
icd9_patients_list.sort(key=length_2_elem, reverse=True)
```

```python
icd9_patients_list.sort(key=lambda x: len(x[1]), reverse=True)
```

```python
heicd9_patients_list.sort)
```

```python
icd9_patients_list[-5:]
```

```python
for c, patient_list in icd9_patients_list:
    print(c)
    for p in patient_list:
        print("\t",p)
```

```python
for c, patient_list in icd9_patients_list:
    print(icd9_map[c], len(patient_list))

```

```python
max([len(v) for v in icd9_map.values()])
```

```Python
print("%s\t%3d"%(icd9_map[c].ljust(60, "."), len(patient_list)))
```
* `%s` means substitute a string into this location of the string
* `\t` tab character
* `%d` means substitute an integer value here
* `%3d` means substitute an integer value here and use at least three spaces
* `%03d` means substitute an integer value here and use at least three spaces left padding with a zero

```python
for c, patient_list in icd9_patients_list:
    print("%s\t%3d"%(icd9_map[c].ljust(60, "."), len(patient_list)))
```

### Repeat with a regular dictionary

```python
icd9_patients = {}
for p,c,l in icd9_data[:50]:
    print(p,c,l)
    if c in icd9_patients:
        icd9_patients[c].append(p)
    else:
        new_list = [p]
        icd9_patients[c] = new_list
    print(icd9_patients)
    input("Continue")
icd9_patients_list=list(icd9_patients.items())
icd9_patients_list.sort(key = lambda x: len(x[1]), reverse=True)
for c, patient_list in icd9_patients_list:
    print("%s\t%3d"%(icd9_map[c].ljust(60, "."), len(patient_list)))
```

```Python

icd9_patients = defaultdict(list)
for p,c,l in icd9_data:
    try:
        icd9_patients[c].append(p)
    except:
        pass

#######################

for p,c,l in icd9_data:
    if c in icd9_patients:
        icd9_patients[c].append(p)
    else:
        new_list = [p]
        icd9_patients[c] = new_list
        
```

```python
d3 = defaultdict(int)
d3["Brian"]
d3["Wendy"] = d3["Wendy"]+29
d3
```

```python
icd9_patients = {}
for p,c,l in icd9_data:
    if c in icd9_patients:
        icd9_patients[c].append(p)
    else:
        new_list = [p]
        icd9_patients[c] = new_list


```

```python
for c, patient_list in icd9_patients.items():
    print("%s\t%3d"%(icd9_map[c].ljust(60, "."), len(patient_list)))
```

```python
from collections import Counter
Counter("Brian Brian Wendy Wendy Wendy Argos Helios Argos")
```

```python
type(("9",))
```

```python
import random
d,p = random.choice(icd9_patients_list)
print(d,len(p),sep="\n")
print(icd9_patients["V12.59"])
for d,p in icd9_patients.items():
    if "HX" in d:
        print(d,len(p))
```

```python
assert_equals(len(icd9_patients["V12.59"]),5)
assert_equals(len(icd9_patients["572.2"]),12)
```

#### Loop through  `icd9_patients_list`
1. for each element in `icd9_patients_list` print the ICD9 label corresponding to the code and the number of patients with each diagnosis.

<!-- grade_id: cell-121dd04abe83faee, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### How many diagnoses does each patient have?

<!-- grade_id: cell-305cd7b053c0116a, points: 0, grade: false, solution: true -->
```python
patient_diagnoses = defaultdict(list)
# YOUR CODE HERE
raise NotImplementedError()
```

```python
for p, d in patient_diagnoses_list:
    print(p)
    print(d) # replace with len(d)
    print("\n")
```

<!-- grade_id: cell-322e571e6bd058c0, points: 0, grade: true, solution: false -->
```python
assert_equals(len(patient_diagnoses[2512]),49)
assert_equals(len(patient_diagnoses[353]),56)
assert_equals(len(patient_diagnoses_list[0][1]),125)
assert_equals(len(patient_diagnoses_list[45][1]),23)

```

```python

```