```julia
Pkg.add("DataFrames")
Pkg.add("Seaborn");
```

```julia
using DataFrames
using Base.Test
using Seaborn;
using Base.Random
```

# Playing with Primes

## Encyrption is in the Center of Public Debate
* [CNN, "ISIS using encryption to evade FBI"](http://www.cnn.com/2015/07/08/politics/fbi-comey-isis-encryption-recruitment/)
* [Wired, "Security Manual Reveals the OPSEC Advice ISIS Gives Recruits"](http://www.wired.com/2015/11/isis-opsec-encryption-manuals-reveal-terrorist-group-security-protocols/)
* [Fortune, "Former National Security Officials Support Encryption"](http://fortune.com/2015/12/16/national-security-encryption/)

## Encryption Relies on Mathematics
### Encryption Relies on Prime Numbers


>[Public-key] was invented in 1976 by two Stanford
mathematicians, Whitfield Diffie and Martin Hellman. Their discovery can be phrased simply:
enciphering schemes should be asymmetric. For thousands of years all ciphers were symmetric –
the key for encrypting a message was identical to the key for decrypting it, but used, so to speak,
in reverse. To change “5 100 100 5 15 55” or “6 120 120 6 18 66” back into “attack,” for instance,
one simply reverses the encryption by dividing the numbers with the key, instead of multiplying
them, and then replaces the numbers with their equivalent letters. Thus sender and receiver must
both have the key, and must both keep it secret. The symmetry, Diffie and Hellman realized, is the
origin of the key-management problem. The solution is to have an encrypting key that is different
from the decrypting key – one key to encipher a message, and another, different key to decipher
it. With an asymmetric cipher, Alice could send encrypted messages to Bob without providing
him with a secret key. In fact, Alice could send him a secret message even if she had never before
communicated with him in any way.

>“If this sounds ridiculous, it should,” Schneier wrote in Secrets and Lies (2001). “It sounds
impossible. If you were to survey the world’s cryptographers in 1975, they would all have told you
it was impossible.” One year later, Diffie and Hellman showed that it was possible, after all. (Later
the British Secret Service revealed that it had invented these techniques before Diffie and Hellman,
but kept them secret – and apparently did nothing with them.)

>To be precise, Diffie and Hellman demonstrated only that public-key encryption was possible in
theory. Another year passed before three MIT mathematicians – Ronald L. Rivest, Adi Shamir,
and Leonard M. Adleman – figured out a way to do it in the real world. At the base of the RivestShamir-Adleman,
or RSA, encryption scheme is the mathematical task of factoring. Factoring
a number means identifying the prime numbers which, when multiplied together, produce that
number. Thus 126,356 can be factored into $2 \times 2 \times 31 \times 1,019$, where 2, 31, and 1,019 are all
prime. (A given number has only one set of prime factors.) 1 Surprisingly, mathematicians
regard factoring numbers – part of the elementary-school curriculum – as a fantastically difficult
task. Despite the efforts of such luminaries as Fermat, Gauss, and Fibonacci, nobody has ever
discovered a consistent, usable method for factoring large numbers. Instead, mathematicians try
potential factors by invoking complex rules of thumb, looking for numbers that divide evenly. For
big numbers the process is horribly time-consuming, even with fast computers. The largest number
yet factored is 155 digits long. It took 292 computers, most of them fast workstations, more than
seven months.

>Note something odd. **It is easy to multiply primes together. But there is no easy way to take
the product and reduce it back to its original primes.** In crypto jargon, this is a “trapdoor”: a
function that lets you go one way easily, but not the other. Such one-way functions, of which
this is perhaps the simplest example, are at the bottom of all public-key encryption. They make
asymmetric ciphers possible.

>To use RSA encryption, Alice first secretly chooses two prime numbers, p and q, each more than
a hundred digits long. This is easier than it may sound: there are an infinite supply of prime
numbers. Last year a Canadian college student found the biggest known prime: $2^{13466917} − 1$. It
has 4,053,946 digits; typed without commas in standard 12-point type, the number would be more
than ten miles long. Fortunately Alice doesn’t need one nearly that big. She runs a program that
randomly selects two prime numbers for her and then she multiplies them by each other, producing
pq, a still bigger number that is, naturally, not prime. This is Alice’s“public key.” (In fact, creating
the key is more complicated than I suggest here, but not wildly so.)

>As the name suggests, public keys are not secret; indeed, the Alices of this world often post
them on the Internet or attach them to the bottom of their e-mail. When Bob wants to send Alice
a secret message, he first converts the text of the message into a number. Perhaps, as before, he
transforms “attack” into “5 100 100 5 15 55.” Then he obtains Alice’s public key – that is, the
number pq – by looking it up on a Web site or copying it from her e-mail. (Note here that Bob does
not use his own key to send Alice a message, as in regular encryption. Instead, he uses Alice’s key.)
Having found Alice’s public key, he plugs it into a special algorithm invented by Rivest, Shamir,
and Adleman to encrypt the message.

>At this point the three mathematicians’ cleverness becomes evident. Bob knows the product pq,
because Alice has displayed it on her Web site. But he almost certainly does not know p and q
themselves, because they are its only factors, and factoring large numbers is effectively impossible.
Yet the algorithm is constructed in such a way that to decipher the message the recipient must
know both p and q individually. Because only Alice knows p and q, Bob can send secret messages to
Alice without ever having to swap keys. Anyone else who wants to read the message will somehow
have to factor pq back into the prime numbers p and q. (["The science of encryption: prime numbers and mod n arithmetic"](https://math.berkeley.edu/~kpmann/encryption.pdf))

## How Do We Find Primes?

A really old way is known as the [Sieve of Eratosthenes](https://en.wikipedia.org/wiki/Sieve_of_Eratosthenes)

The approach here is to create an array of numbers and cross out any number that is a multiple of a prime. As we loop through the array, any number that is not crossed out is a prime.

1. Allocate an array of integers (`nums`) between 2 and N+1
1. Start with i=0
1. Look at the ith number in the array (`nums[i]`), if it is not crossed out it is a prime.
    1. Add `nums[i]` to our list of primes
    1. Cross out all the multiples of `nums[i]`
    1. Increment i

```julia
a = collect(2:20)
a[3:2:end]=-1
print(a)
```

<!-- grade_id: cell-906b646f7da82489, points: 0, grade: false, solution: true -->
```julia
function sievePrime(N)
    primes = []
    nums = collect(2:N)
    for i in range(1,length(nums))
        current_num = nums[i]
        if current_num != -1
            push!(primes, current_num)
            nums[i+current_num:current_num:end] = -1
        end
    end
    return primes
end
function sievePrime(N)
    nums = collect(2:N)
    for i in range(1,length(nums))
        current_num = nums[i]
        if current_num != -1
            nums[i+current_num:current_num:end] = -1
        end
    end
    return [n for n in nums if n != -1]
end
```

```julia
function isprime(n)
    return n in sievePrime(n)
end
```

```julia
isprime(23)
```

<!-- grade_id: cell-87d7b7740776091c, points: 0, grade: true, solution: false -->
```julia
@test sievePrime(10) == [2, 3, 5, 7]
@test sievePrime(20) == [2, 3, 5, 7, 11, 13, 17, 19]
```

## What are some limitations of this approach?
1. We have to allocate a potentially large array (memory)
1. We loop through the entire list of numbers

### Given `N`, what is the largest number we would have to test division on to see if `N` is prime? 

### Let's write some alternative functions for finding primes

## Exercise:

modify `isprime2` to check whether a number is prime. Use `if` statements to test for the special cases ($n<2$, $n=2$, $n\mod2=0$). Use a for loop and the range function ($\text{range}(3,1+\sqrt{n}, 2)$)? Why $3$? Why $2$?

<!-- grade_id: cell-839df251964424e3, points: 0, grade: false, solution: true -->
```julia
function isprime2(n)
    if !isa(n,Int)
        throws(ArgumentError("$n: n must be an integer"))
    end
    if n < 2
        return false
    elseif n == 2
        return true
    elseif n %2 == 0
        return false
    else
        lim = 1+convert(Int64, round(sqrt(n)))
        for i in collect(3:2:lim)
            if n % i == 0
                return false
            end
        end
    end
    return true
end
```

<!-- grade_id: cell-d7f0ae1402b7514b, points: 0, grade: true, solution: false -->
```julia
@test isprime2(47)
@test !isprime2(49)
```

## Exercise

Modify `isprime3` to check if a number `p` is Prime using list comprehension and the [`any`](https://docs.python.org/3/library/functions.html#any) function.

<!-- grade_id: cell-701c0e3f0b610db8, points: 0, grade: false, solution: true -->
```julia
function isprime3(p)
    if !isa(p,Int)
        throws(ArgumentError("$p: p must be an integer"))
    end
    if p < 2
        return false
    elseif p == 2
        return true
    elseif p % 2 == 0
        return false
    else
        return  ! any(p%n==0 for n in collect(3:2:convert(Int64, round(sqrt(p)))+1)) 
    end
end
```

```julia
@test !isprime3(49)
@test isprime3(47);

```

```julia
primes = [n for n in range(1,1000) if isprime2(n)]
primes
```

```julia
@time [n for n in collect(1:100000) if isprime(n)];
```

```julia
@time [n for n in collect(1:100000) if isprime2(n)];
```

```julia
@time [n for n in range(1,100000) if isprime3(n)];
```

## [Gilbreath's Conjecture](https://en.wikipedia.org/wiki/Gilbreath%27s_conjecture)

>Gilbreath's conjecture is a conjecture in number theory regarding the sequences generated by applying the forward difference operator to consecutive prime numbers and leaving the results unsigned, and then repeating this process on consecutive terms in the resulting sequence, and so forth. The statement is named after mathematician Norman L. Gilbreath who, in 1958, presented it to the mathematical community after observing the pattern by chance while doing arithmetic on a napkin. In 1878, eighty years before Gilbreath's discovery, François Proth had, however, published the same observations along with an attempted proof, which was later shown to be false.

#### Write a function ``seq_diff`` that computes the absolute value of the forward difference of a sequence of primes

**Hints:**
1. The `math` library defines an `fabs` which returns a floating point value for the absolute value.
1. Use the `range` function.
    1. What do you want for the starting value for `range`?
    1. What do you want for the ending value for `range`?
1. How will the length of the output differ from the length of the input?

<!-- grade_id: cell-db3bfe56e4780385, points: 0, grade: false, solution: true -->
```julia
function seq_diff(seq)
    return [abs(seq[i]-seq[i-1]) for i in collect(2:length(seq))]
end

```

#### Test the output

<!-- grade_id: cell-686cd642a626638c, points: 0, grade: true, solution: false -->
```julia
seq = [2, 3, 5, 7, 11, 13, 17, 19, 23, 29, 31, 37, 41, 43, 47]
assert_equal(seq_diff(seq), [1, 2, 2, 4, 2, 4, 2, 4, 6, 2, 6, 4, 2, 4])
assert_equal(seq_diff(seq_diff(seq)), [1, 0, 2, 2, 2, 2, 2, 2, 4, 4, 2, 2, 2])
```

```julia
seq = [n for n in range(1,50) if isprime2(n)]
println(seq)
while true
    if isempty(seq)
        break   
    end
    seq = seq_diff(seq)
    println(seq)
end
```

## Playing with Large Measured Data: [Benford's Law](https://en.wikipedia.org/wiki/Benford%27s_law)

### We are going to use census data

```julia
data = readtable("../InClass/PEP_2016_PEPANNRSIP.US12A_with_ann.csv",
skipstart=1)
```

```julia
city_populations = readtable("../InClass/PEP_2016_PEPANNRSIP.US12A_with_ann.csv",
skipstart=1)[17];
```

## `city_populations` is a list of integers that are the populations of USA cities

### Create a list of all the populations greater than or equal to 100000 where the integers are converted to strings

```julia
rand(5)
```

<!-- grade_id: cell-0324eead7bbbaddc, points: 0, grade: false, solution: true -->
```julia
### BEGIN SOLUTION
function get_string_nums(nums; thresh=99999)
    return [string(x) for x in nums if x > thresh]
end
### END SOLUTION
```

```julia
get_first_digits(get_string_nums(city_populations))
```

### Create a list named `city_populations_first_digits` consisting of the first digit in each number converted to an int

```julia
for x in city_populations_strings
    print(parse(Int64, x[1]))
end
```

<!-- grade_id: cell-007a102e12198872, points: 0, grade: false, solution: true -->
```julia
function get_first_digits(nums)
    return [ parse(Int64, x[1]) for x in nums];
end
```

```julia
city_populations_first_digits = get_first_digits(get_string_nums(city_populations));
```

### What does the distribution of first digits look like?

#### It should look like this
![digit distribution](./population_digits.png)

```julia
distplot(city_populations_first_digits, kde=false)
savefig("population_digits.png")
```

## How does this differ from randomly generated numbers?

### We have to duplicate all our prior code...
### ...Unless we put all the functionality in functions.

The resulting distribution should look something this:

![random digits](./random_digits.png)

```julia
rn = rand(99999:10000000, length(city_populations_strings));
```

```julia
distplot(get_first_digits(get_string_nums(rn)), kde=false)

```

```julia

```