# Design Decisions with Python Functions



```python
from nose.tools import assert_true, assert_equal, assert_false, assert_almost_equal, assert_raises
```

# Two primary reasons for defining functions:
1. Code reuse: 
    * Write and debug code once. Then I can use this same (correct) code many times
    * This makes upkeep/modifications simpler. When I think of an improved way of implementing something I only need to change it in one location.

2. Procedural Decomposition:
    * A function should do one thing, not multiple things.
    * This can become a matter of style
    

## Function Style According to Mark Thomason
![Mark Thomason](./mark_thomason.jpg)

## The entire function should be visible on your screen within your editor.
## If your function doesn't fit on your screen, get a bigger screen

### What are the implications of these heuristic?
#### Function size changes with age?

## Exercise: Define a function to get a positive integer from a user.
### Requirements
1. Use an infinite while loop
1. Use the input function
1. Keep prompting the user for input until a valid positive integer is provided

We'll use a try/except block to account for users entering non-integer value

```Python
try:
    # Get input from user
    # convert to an integer (this where we could get an exception
    # test for positivity
except ValueError:
    # if we get an input that we can't convert to an integer, we need to do something
```

```python
def get_pos_integer(prompt="Enter a positive integer"):
    while True:
        num = input(prompt)
        try:
            num = int(num)
            if num > 0:
                return num
        except ValueError:
            pass

```

```python
get_pos_integer()
```

## Does this function do one thing?
## Could we break it into smaller pieces?

### Write a function to test whether a number is positive

<!-- grade_id: cell-b17130446493e1e7, points: 0, grade: false, solution: true -->
```python
def ispositive(x):
    return x > 0
```

<!-- grade_id: cell-163d61c8d4528ae4, points: 0, grade: true, solution: false -->
```python
assert_true(ispositive(5))
assert_false(ispositive(-1))
assert_false(ispositive(0))
```

<!-- grade_id: cell-2514734b675a7247, points: 0, grade: false, solution: true -->
```python
def getint(sint):
    return int(sint)
    
```

```python
assert_equal(getint("543"), 543)
assert_raises(ValueError, getint, "4.7")
getint("4.7")
```

<!-- grade_id: cell-c1513c8de2dfd7da, points: 0, grade: true, solution: false -->
```python
assert_equal(getint(543), 543)
assert_raises(ValueError, getint, "4.7")

```

### Using `getint` and `ispositive` rewrite `get_pos_integer`

<!-- grade_id: cell-fa342b9026507d73, points: 0, grade: false, solution: true -->
```python
def get_pos_integer2(prompt="enter a positive integer: "):
    while True:
        num = input(prompt)
        try:
            num = getint(num)
            if ispositive(num):
                return num
        except ValueError:
            pass
```

<!-- grade_id: cell-e69024ff74b045be, points: 0, grade: true, solution: false -->
```python
# Enter '5'
assert_equal(get_pos_integer2(),5)

# Enter '7'
assert_equal(get_pos_integer2(),7)

```

## Exercise

Following the same style as `get_pos_integer`, write a function `get_value`. `get_value` takes as arguments:

1. A positional argument `converter` that is a function that takes as input a string and returns the desired value
1. A positional argument `tester` that is a function that takes as input a value and returns `True` or `False` depending on whether a desired condition is satisfied.
1. A keyword argument `prompt` that is the prompt to use with `input`.

Test the function with `getint` and `ispositive`.

<!-- grade_id: cell-cebe2a0d20841727, points: 0, grade: false, solution: true -->
```python
def get_value(converter, tester, prompt="enter a positive integer"):
    while True:
        num = input(prompt)
        try:
            num = converter(num)
            if tester(num):
                return num
        except ValueError:
            pass    
```

```python
get_value(getint, ispositive)
```

## Exercise

Modify `get_three_words` to take a string `entry` splits it on white spaces and returns a list of three words. If the list is not three words long, raise a `ValueError`.

<!-- grade_id: cell-ade3a80c6dfc2cd7, points: 0, grade: false, solution: true -->
```python
def get_three_words(entry):
    # YOUR CODE HERE
    words = entry.split()
    if len(words) == 3:
        return words
    else:
        raise ValueError("You need to enter three words separarted by spaces")
```

```python
int("5.7")
```

```python
get_three_words("Brian, Earl, Chapman")
```

## Exercise

Modify `test_ascending` to take in a sequence and test if the elements in `values` are in ascending order.

**Hint:** Use the `all` function.

<!-- grade_id: cell-40b8ef1cd8e54c0b, points: 0, grade: false, solution: true -->
```python
def test_ascending(values):
    for i in range(len(values)-1):
        if not values[i] < values[i+1]:
            return False
    return True

def test_ascending(values):
    return all([values[i] < values[i+1] for i in range(len(values)-1)])
```

```Python
test_ascending((1,2,3)) == True
test_ascending((3,2,1)) == False
test_ascending(("Argos", "Helios", "Zeus"))==True
```

<!-- grade_id: cell-519436fcb56d35ad, points: 0, grade: true, solution: false -->
```python
assert_true(test_ascending(("Argos", "Helios", "Zeus")))
assert_false(test_ascending(("Argos", "Zeus", "Helios")))
assert_false(test_ascending(("argos", "Helios", "Zeus")))

```

```python
get_value(get_three_words, 
          test_ascending, 
          prompt="enter three words in ascending alphabetical order separated by spaces: ")

```

## [Avoiding empty list as a default argument](https://stackoverflow.com/questions/1132941/least-astonishment-and-the-mutable-default-argument)

## Observe this unexpected behavior

```python
def foo(a=[]):
    a.append(5)
    return a
print(foo())
print(foo())
```

```python
def i_hate_foo(a=None):
    if a == None:
        a = []
    a.append(5)
    return a
print(i_hate_foo())
print(i_hate_foo())
```

```python
i_hate_foo([3,4])
i_hate_foo()
```

### Variable number of position arguments

A function definition that looks like this

```python

def some_function(*x):
    ### BLOCK OF CODE
```

Has a variable number of positional arguments.

The variable number of positional arguments are passed to the function as a **list.**

```python
def demo1(*x):
    for xx in x:
        print(xx)
demo1(1,"Brian", [1,2,3], {4,5,6})
print("*"*42)
demo1(1,2,3)
print("*"*42)

demo1()
```

## Exercise

Write a function ``sumthings`` that takes a variable number of arguments and returns the sum of their values.

#### Challenge: Can you do this with a single Python statement within `sumthings`?

```python
def sumthings(*x):
    value = 0
    for xx in x:
        value += xx
    return value
```

<!-- grade_id: cell-e7d16b80330ab587, points: 0, grade: false, solution: true -->
```python
def sumthings(*x):
    return sum(x)
```

<!-- grade_id: cell-28f6eb94cc2d3bbc, points: 0, grade: true, solution: false -->
```python
assert_equal(sumthings(1,2,3),6)
assert_equal(sumthings(1),1)
assert_equal(sumthings(),0)

```

### Variable number of keyword arguments

A function definition that looks like this

```python

def some_function(**kwargs):
    ### BLOCK OF CODE
```

has a variable number of keyword arguments

The variable number of positional arguments are passed to the function as a **list.**

```python
def demo2(**kwargs):
    for k in kwargs:
        print(k, kwargs[k])
demo2(print="No way", age=29, favorite_number=8, luck="bad")
print("*"*42)
demo2(sings=False, runs=False, bikes=True)
```

```python
def demo3(*x, keyword=None, **kwargs):
    print(keyword)
demo3(1,2,3,name="Brian")
```

```python
def demo3(keyword=None):
    print(keyword)
demo3(keyword="Brian")
```

## Exercise

Write a function `keep_numeric` that takes a variable number of keyword arguments and returns a dictionary consisting of the kwargs passed to `keep_numeric` where the value is an integer

**Challenge**: Return a dictionary of kwargs where the value is any numeric value

<!-- grade_id: cell-02479618f82ab512, points: 0, grade: false, solution: true -->
```python
import numbers
def keep_numeric(**kwargs):
    value = {}
    for k, v in kwargs.items():
        if type(v) ==int or  type(v) == float:
            value[k] = v
    return value
    
def keep_numeric(**kwargs):
    value = {}
    for k, v in kwargs.items():
        if isinstance(v, numbers.Number):
            value[k] = v
    return value
    
```

<!-- grade_id: cell-1539513e821c8af2, points: 0, grade: true, solution: false -->
```python
assert_equal(keep_numeric(name="Brian", age=29, race="Caucasian", weight=150), {'age': 29, 'weight': 150})
```

<!-- grade_id: cell-73691bdbd54bbd04, points: 0, grade: true, solution: false -->
```python
assert_equal(keep_numeric(name="Brian", age=29.5, race="Caucasian", weight=150.4), {'age': 29.5, 'weight': 150.4})
```

<!-- grade_id: cell-eeb2dc429e508bc3, points: 0, grade: true, solution: false -->
```python
assert_equal(keep_numeric(name="Brian", age=29.5, race="Caucasian", weight=150.4), {'age': 29.5, 'weight': 150.4})
```

```python
assert_equal(keep_numeric(name="Brian", age=29.5+0j, race="Caucasian", weight=150.4), {'age': 29.5, 'weight': 150.4})
```

```python

```