# BMI 6018
## Fall 2017
### M3-Numerics Homework 

```python
import numpy as np
%matplotlib inline
from nose.tools import assert_equal
from numpy.testing import assert_array_equal, assert_allclose
import seaborn as sns
```

### Additional packages

We are going to install some additional packages using the *conda* function that comes with Anconda.

* ``ipywidgets`` will allow us to use interactive widgets in the notebook
* ``holoviews`` is a visualization package that allows you to attach semantic meaning to data
* ``seaborn`` is an extension/improvement of matplotlib

**Problem 1 (20 points):** BMI is a person's weight in kilograms divided by the square of height in meters. 

Below I am reading in actual heights (in inches) and weights (in kgs) of individual patients in the MIMIC2 database. 

Modify the function `bmi` to return an array with the bmi value computed for each person given their height and weight. The elements of the arrays are aligned so the nth element of each array corresponds to the nth patient.

```python
h = np.fromfile("./numeric_mimic_heights.txt",sep=",")
w = np.fromfile("./numeric_mimic_weights.txt",sep=",")
```

<!-- grade_id: cell-fed21420b6e4c717, points: 0, grade: false, solution: true -->
```python
def bmi(h, w):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-8ae1df262410562f, points: 20, grade: true, solution: false -->
```python
### BEGIN HIDDEN TEST
ht = np.random.normal(60, 10, size=100)
wt = np.random.normal(70, 20, size=100)
assert_array_equal(bmi(ht, wt), 
                   wt/(ht*2.54/100.0)**2)
### END HIDDEN TEST
```

### Here is a visualizations  of the data 

We'll learn more details about these packages later in the class.

```python
sns.distplot(bmi(h,w),color="m")

```

**Problem 2 (20 points):** Finding roots of a quadratic equation.

For the general quadratic equation $ax^2+bx+c=0$, the quadratic formula is $x=\frac{-b\pm\sqrt{b^2-4ac}}{2a}$.

I have provided a skeleton code below for a function that takes three arguments (a,b,c) and computes the the roots for the quadratic equation with these coefficients. Using the [math library](https://docs.python.org/3/library/math.html), complete the function `find_roots` to return the two roots for $x$ given $a, b, c$. 

<!-- grade_id: cell-ace06b2ebf6ed375, points: 0, grade: false, solution: true -->
```python
import math

def find_roots(a=1,b=1,c=1):
    """
    This is a docstring. It is what the help() function returns
    about this function.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6d30fbe175359fae, points: 20, grade: true, solution: false -->
```python
assert_equal(find_roots(1,5,6),(-3.0, -2.0))
assert_equal(find_roots(1,-1,-2), (-1.0, 2.0))
assert_equal(find_roots(2,4,-4), (-2.732050807568877, 0.7320508075688772))
```

**Problem 3 (10 Points):** HIPAA allows researchers who do not have informed consent from study subjects to only know their age within decades. Modify the function `age_string` to take as an input argument an age in years (numeric) and returns a string giving the age in decades (e.g. 27 &#8594; "Age in 20s").

Note in my code fragment below, I use c-style string formatting to insert the computer program that takes an age in years and return a HIPAA compliant description of the age.

<!-- grade_id: cell-576fc1ef45177db5, points: 0, grade: false, solution: true -->
```python
def age_string(age):
    decade = None
    # YOUR CODE HERE
    raise NotImplementedError()
    return "Age in %ds"%decade
    

```

<!-- grade_id: cell-cae5746e4897cc18, points: 10, grade: true, solution: false -->
```python
assert_equal(age_string(28.9), "Age in 20s")
assert_equal(age_string(55), "Age in 50s")
assert_equal(age_string(6), "Age in 0s")
```

**Problem 4 (20 points):** X-ray attenuation in the human body can be described with the Lambert-Beers law 
$$ I = I_0e^{-\mu L}$$
where $I_0$ is the incident x-ray intensity, I is the transmitted x-ray intensity, *µ* is the linear attenuation coefficient of the material, and *L* is the thickness of the material. µ is a combination of the attenuation due to photoelectric interactions, Compton interactions, and Rayleigh scattering. In medical applications, the linear attenuation is usually reported as the Hounsfield unit (HU) defined as 
$$HU = 1000\times \frac{\mu-\mu_{\text{water}}}{\mu_{\text{water}}-\mu{\text{air}}},$$ 
where $\mu_{\text{water}} =0.1928 \text{cm}^{-1}$ and $\mu_{\text{air}} = 0.0002 \text{cm}^{-1}$. 
    
Modify the function `hounsfield` that takes a transmitted x-ray intensity, an incident x-ray intensity, and object thickness and returns the average HU for the object.

<!-- grade_id: cell-54635d02f4cdcc23, points: 0, grade: false, solution: true -->
```python

def hounsfield(I, I0, L):
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-bd6a7a42e26f0fe3, points: 20, grade: true, solution: false -->
```python
assert_equal(hounsfield(0.2, 1.0, 8), 0.20117973905426254)
assert_equal(hounsfield(0.4, 1.0, 4), 0.22907268296853875)
assert_equal(hounsfield(0.8, 1.0, 10), 0.022314355131420972)

```

**Problem 5 (10 points):** Thick of a problem you might be interested in tackling for your term project. Describe numeric data that could be associated with your term project. Would you represent the data with floats or integers?

<!-- grade_id: cell-48f88cf9f0ab8fa6, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

**Problem 6 (10 points):** In the problems below we denote an intger in some arbitrary base $b$ as $xyz_b$. For each of the problems below, convert the numbers to base 10 integers. Change the solution integers (e.g. `n1`) from `None` to the appropriate base 10 integer (numeric value).

$101101_2$ (Binary number, base-2)

<!-- grade_id: cell-8ff97c44e81be2ae, points: 0, grade: false, solution: true -->
```python
n1 = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6820522b2ae2095d, points: 2, grade: true, solution: false -->
```python

```

$D_{16}$ (Base 16)

<!-- grade_id: cell-d9b65419b3daee61, points: 0, grade: false, solution: true -->
```python
n2 = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-810f6a81f4743c56, points: 0, grade: true, solution: false -->
```python

```

$120_7$

<!-- grade_id: cell-33196d2411e3185c, points: 0, grade: false, solution: true -->
```python
n3 = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-10c0a016ed319d90, points: 2, grade: true, solution: false -->
```python

```

$12_3$

<!-- grade_id: cell-89db955579d3353c, points: 0, grade: false, solution: true -->
```python
n4 = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-673b95e7b12fec5b, points: 2, grade: true, solution: false -->
```python

```

$31_8$

<!-- grade_id: cell-ec883a833c4ad16b, points: 0, grade: false, solution: true -->
```python
n5 = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5f2d141829d314aa, points: 2, grade: true, solution: false -->
```python

```